// Numbas version: exam_results_page_options
{"name": "Teste", "metadata": {"description": "", "licence": "Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International"}, "duration": 0, "percentPass": 0, "showQuestionGroupNames": false, "showstudentname": true, "question_groups": [{"name": "Group", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": [""], "questions": [{"name": "Equazione esplicita di una retta dal grafico e della sua parallela e perpendicolare per un punto dato", "extensions": ["geogebra"], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Mark Hodds", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/510/"}, {"name": "Irene Paniccia", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/4869/"}], "tags": [], "metadata": {"description": "<p style=\"-webkit-user-select: auto;\">Utilizza un grafico di una retta $y=mx+c$ incluso da Geogebra con coefficienti casuali impostati da NUMBAS. Il grafico contiene anche un punto $P$ di coordinate casuali impostate da NUMBAS. L'esercizio chiede di determinare l'equazione della retta e quelle delle rette parallela e perpendicolare per $P$.</p>", "licence": "Creative Commons Attribution 4.0 International"}, "statement": "<p>Il&nbsp;grafico che segue mostra la&nbsp;retta $r$ nel piano cartesiano e il punto $P$ a essa esterno.</p>\n<p>{geogebra_applet('h46bjwqv', defs, [])}</p>", "advice": "<p style=\"text-align: justify;\"><strong>a)</strong> Per determinare l'equazione della retta $r$,&nbsp;osserviamo&nbsp;innanzitutto che, per ogni avanzamento di un'unit&agrave; verso destra, la retta avanza di $\\simplify[fractionNumbers]{{m}}$ in verticale. Dunque $m = \\simplify[fractionNumbers]{{m}}$.&nbsp;Essendo l'intersezione con l'asse $y$ in $( 0 ; \\var{q})$, si ha $q = \\var{q}$.</p>\n<p style=\"text-align: justify;\">In definitiva</p>\n<p style=\"text-align: justify;\">\\begin{equation} r : \\, y = \\simplify[zeroFactor,fractionNumbers,unitFactor]{{m} x+ {q}} . \\end{equation}</p>\n<p style=\"text-align: justify;\"><strong>b)</strong> Per determinare le coordinate del punto $P$, &egrave; sufficiente osservare il grafico: si ha $P = ({x_P};{y_P})$.</p>\n<p style=\"text-align: justify;\"><strong>c)</strong>&nbsp;Il fascio (improprio) di rette parallele alla retta data&nbsp;ha equazione&nbsp;$y = \\simplify[zeroFactor,fractionNumbers,unitFactor]{{m} x + q}$.&nbsp;Imponendo il passaggio di una retta del fascio per il punto $P$ otteniamo $q = \\simplify[zeroFactor,fractionNumbers,unitFactor]{{y_P - m*x_P}}$.</p>\n<p style=\"text-align: justify;\">Pertanto</p>\n<p style=\"text-align: justify;\">\\begin{equation} s : \\, y = \\simplify[zeroFactor,fractionNumbers,unitFactor]{{m} x + {y_P - m*x_P}} . \\end{equation}</p>\n<p style=\"text-align: justify;\"><strong>d)</strong> Il fascio (improprio) di rette perpendicolari alla retta data ha equazione $y = \\simplify[zeroFactor,fractionNumbers,unitFactor]{{- 1/m} x + q}$.&nbsp;Imponendo il passaggio di una retta del fascio per il punto $P$ otteniamo $q = \\simplify[zeroFactor,fractionNumbers,unitFactor]{{y_P + 1/m*x_P}}$.</p>\n<p style=\"text-align: justify;\">Pertanto</p>\n<p style=\"text-align: justify;\">\\begin{equation} t : \\, y = \\simplify[zeroFactor,fractionNumbers,unitFactor]{{-1/m} x + {y_P +1/m*x_P}} . \\end{equation}</p>", "rulesets": {}, "variables": {"m0": {"name": "m0", "group": "Ungrouped variables", "definition": "(random(random(1..3),random(1/2,2/3,1/4,3/4,1/3)))*random(-1,1)", "description": "<p>Coefficiente angolare della retta r</p>", "templateType": "anything"}, "defs": {"name": "defs", "group": "Ungrouped variables", "definition": "[\n  ['m',m],\n  ['q',q],\n  ['P',P]\n]", "description": "", "templateType": "anything"}, "q0": {"name": "q0", "group": "Ungrouped variables", "definition": "random(-3..3)", "description": "<p>Termine noto (quota) della retta r</p>", "templateType": "anything"}, "P": {"name": "P", "group": "Ungrouped variables", "definition": "vector(x_P,y_P)", "description": "<p>Punto del piano di coordinate intere esterno alla retta r</p>", "templateType": "anything"}, "x_P": {"name": "x_P", "group": "Ungrouped variables", "definition": "random(-6..6 except 0)", "description": "", "templateType": "anything"}, "y_P": {"name": "y_P", "group": "Ungrouped variables", "definition": "random(-4..4 except m0*x_P+q0)", "description": "", "templateType": "anything"}, "m": {"name": "m", "group": "Ungrouped variables", "definition": "m0", "description": "", "templateType": "anything"}, "q": {"name": "q", "group": "Ungrouped variables", "definition": "q0", "description": "", "templateType": "anything"}}, "variablesTest": {"condition": "", "maxRuns": "200"}, "ungrouped_variables": ["defs", "m0", "q0", "P", "x_P", "y_P", "m", "q"], "variable_groups": [], "functions": {"graph2": {"parameters": [], "type": "html", "language": "jme", "definition": ""}, "graph": {"parameters": [], "type": "html", "language": "jme", "definition": ""}, "graph3": {"parameters": [], "type": "html", "language": "jme", "definition": ""}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": true, "customName": "a)", "marks": 0, "showCorrectAnswer": true, "showFeedbackIcon": false, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "prompt": "<p>Qual &egrave; l'equazione della retta $r$?</p>\n<p>$r : \\, y = $&nbsp;[[0]]</p>", "gaps": [{"type": "jme", "useCustomName": true, "customName": "equazione", "marks": "1", "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "answer": "{m}x+{q}", "answerSimplification": "zeroFactor,fractionNumbers,unitFactor", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "valuegenerators": [{"name": "x", "value": ""}]}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "b)", "marks": 0, "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "prompt": "<p>Quali sono le coordinate del punto $P$?</p>\n<p>$P = $ ([[0]];[[1]])</p>", "gaps": [{"type": "numberentry", "useCustomName": true, "customName": "x_P", "marks": "0.5", "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "minValue": "x_P", "maxValue": "x_P", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "si-en", "si-fr", "plain-eu"], "correctAnswerStyle": "plain-eu"}, {"type": "numberentry", "useCustomName": true, "customName": "y_P", "marks": "0.5", "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "minValue": "y_P", "maxValue": "y_P", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "si-en", "si-fr", "plain-eu"], "correctAnswerStyle": "plain-eu"}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "c)", "marks": 0, "showCorrectAnswer": true, "showFeedbackIcon": false, "scripts": {}, "variableReplacements": [{"variable": "m", "part": "p0", "must_go_first": false}, {"variable": "q", "part": "p0", "must_go_first": false}], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "prompt": "<p>Determina&nbsp;l'equazione della retta $s$, parallela a $r$ per il punto $P$</p>\n<p>$s: \\, y = $&nbsp;[[0]]</p>", "gaps": [{"type": "jme", "useCustomName": true, "customName": "parallela", "marks": "1", "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "answer": "{m}x+{y_P-m*x_P}", "answerSimplification": "zeroFactor,fractionNumbers,unitFactor", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "valuegenerators": [{"name": "x", "value": ""}]}], "sortAnswers": false}, {"type": "gapfill", "useCustomName": true, "customName": "d)", "marks": 0, "showCorrectAnswer": true, "showFeedbackIcon": false, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "prompt": "<p>Determina&nbsp;l'equazione della retta $t$, perpendicolare&nbsp;a $r$ per il punto $P$</p>\n<p>$t: \\, y = $&nbsp;[[0]]</p>", "gaps": [{"type": "jme", "useCustomName": true, "customName": "perpendicolare", "marks": "1", "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "answer": "{-1/m}x+{y_P+1/m*x_P}", "answerSimplification": "zeroFactor,fractionNumbers,unitFactor", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "valuegenerators": [{"name": "x", "value": ""}]}], "sortAnswers": false}]}]}], "navigation": {"allowregen": false, "reverse": true, "browse": true, "allowsteps": true, "showfrontpage": true, "showresultspage": "oncompletion", "navigatemode": "sequence", "onleave": {"action": "none", "message": ""}, "preventleave": true, "startpassword": ""}, "timing": {"allowPause": true, "timeout": {"action": "none", "message": ""}, "timedwarning": {"action": "none", "message": ""}}, "feedback": {"showactualmark": true, "showtotalmark": true, "showanswerstate": true, "allowrevealanswer": true, "advicethreshold": 0, "intro": "", "reviewshowscore": true, "reviewshowfeedback": true, "reviewshowexpectedanswer": true, "reviewshowadvice": true, "feedbackmessages": []}, "contributors": [{"name": "Fernando Naufel", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/9267/"}], "extensions": ["/srv/numbas/media/user-extensions/extracted/41/geogebra"], "custom_part_types": [], "resources": []}